# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE
    df = pd.read_csv('assets/NISPUF17.csv', usecols=['SEQNUMC', 'EDUC1'])
    
    edu_dic = {"less than high school": 0,
            "high school": 0,
            "more than high school but not college": 0,
             "college": 0}

    df_mask_one = df['EDUC1'] == 1
    edu_dic["less than high school"] = len(df[df_mask_one])
    
    df_mask_two = df['EDUC1'] == 2
    edu_dic["high school"] = len(df[df_mask_two])
    
    df_mask_three = df['EDUC1'] == 3
    edu_dic["more than high school but not college"] = len(df[df_mask_three])
    
    df_mask_four = df['EDUC1'] == 4
    edu_dic["college"] = len(df[df_mask_four])
    
    edu_dic_per = {}
    edu_dic_per["less than high school"] = edu_dic["less than high school"]/len(df)
    edu_dic_per["high school"] = edu_dic["high school"]/len(df)
    edu_dic_per["more than high school but not college"] = edu_dic["more than high school but not college"]/len(df)
    edu_dic_per["college"] = edu_dic["college"]/len(df)
    
    return edu_dic_per
    raise NotImplementedError()
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def average_influenza_doses():
    # YOUR CODE HERE
    df_bf = pd.read_csv('assets/NISPUF17.csv', usecols=['SEQNUMC', 'BF_ENDR06', 'CBF_01', 'P_NUMFLU', 'BF_ENDR06'])
    df_bf.dropna()
    
    df_mask_one = df_bf['CBF_01'] == 1  # breastfed babies
    df_breast = df_bf[df_mask_one]
    df_mask_vax = df_breast['P_NUMFLU'] >= 0  # at least one vaccine
    bf_vax_sum = df_breast[df_mask_vax]
    
    df_mask_two = df_bf['CBF_01'] == 2  # non-breastfed babies
    df_nbreast = df_bf[df_mask_two]
    df_mask_vax_two = df_nbreast['P_NUMFLU'] >= 0 # at least one vaccine
    nbf_vax_sum = df_nbreast[df_mask_vax_two]
    
    bf_avg_vax = df_breast[df_mask_vax]['P_NUMFLU'].mean()
    nbf_vax_avg = df_nbreast[df_mask_vax_two]['P_NUMFLU'].mean()
    
    fluvax_avgs = (bf_avg_vax, nbf_vax_avg)
    print(fluvax_avgs)
    
    return fluvax_avgs
    raise NotImplementedError()
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

Output:
```
(1.8799187420058687, 1.5963945918878317)
```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def chickenpox_by_sex():
    # YOUR CODE HERE
    df_cpox = pd.read_csv('assets/NISPUF17.csv', usecols=['SEQNUMC', 'HAD_CPOX', 'P_NUMVRC', 'SEX'])
    male_vax_stats = {'m_vaxxed': 0, 'm_vaxxed_wpox': 0, 'm_vaxxed_wopox': 0}
    female_vax_stats = {'f_vaxxed': 0, 'f_vaxxed_wpox': 0, 'f_vaxxed_wopox': 0}
    
    pox_rates_by_sex ={}
    
    df_cpox_m = df_cpox['SEX'] == 1  # create the boolean mask
    df_cpox_males = df_cpox[df_cpox_m]  # apply the boolean mask
    df_cpox_vm = df_cpox_males['P_NUMVRC'] >= 1 # create additional boolean mask
    df_cpox_vmales = df_cpox_males[df_cpox_vm]  # apply the boolean mask
    df_cpox_vmyes_mask = df_cpox_vmales['HAD_CPOX'] == 1
    df_cpox_vmyes = df_cpox_vmales[df_cpox_vmyes_mask]
    df_cpox_vmno_mask = df_cpox_vmales['HAD_CPOX'] == 2
    df_cpox_vmno = df_cpox_vmales[df_cpox_vmno_mask]
    
    male_vax_stats['m_vaxxed'] = len(df_cpox_vmales)
    male_vax_stats['m_vaxxed_wpox'] = len(df_cpox_vmyes)
    male_vax_stats['m_vaxxed_wopox'] = len(df_cpox_vmno)
    
    df_cpox_f = df_cpox['SEX'] == 2  # create the boolean mask
    df_cpox_females = df_cpox[df_cpox_f]  # apply the boolean mask
    df_cpox_vfm = df_cpox_females['P_NUMVRC'] >= 1 # create additional boolean mask
    df_cpox_vfemales = df_cpox_females[df_cpox_vfm]  # apply the boolean mask
    df_cpox_vfmyes_mask = df_cpox_vfemales['HAD_CPOX'] == 1
    df_cpox_vfmyes = df_cpox_vfemales[df_cpox_vfmyes_mask]
    df_cpox_vfno_mask = df_cpox_vfemales['HAD_CPOX'] == 2
    df_cpox_vfno = df_cpox_vfemales[df_cpox_vfno_mask]
    
    female_vax_stats['f_vaxxed'] = len(df_cpox_vfemales)
    female_vax_stats['f_vaxxed_wpox'] = len(df_cpox_vfmyes)
    female_vax_stats['f_vaxxed_wopox'] = len(df_cpox_vfno)
    
    male_ratio = male_vax_stats['m_vaxxed_wpox'] / male_vax_stats['m_vaxxed_wopox']
    female_ratio = female_vax_stats['f_vaxxed_wpox'] / female_vax_stats['f_vaxxed_wopox']
    
    pox_rates_by_sex["male"] = male_ratio
    pox_rates_by_sex["female"] = female_ratio
    
    #print(pox_rates_by_sex)
    return pox_rates_by_sex
    raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    '''df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
                   "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])'''
    
    # just return the correlation
    #return corr

    # YOUR CODE HERE
    df_cpox = pd.read_csv('assets/NISPUF17.csv', usecols=['SEQNUMC', 'HAD_CPOX', 'P_NUMVRC'])
    had_pox = []
    num_vax = []
    
    df_wpox_mask = df_cpox['HAD_CPOX'] == 1
    df_wpox = df_cpox[df_wpox_mask]     # children with pox
    df_wpox_vax_mask = df_wpox['P_NUMVRC'] >= 0
    df_wpox_vax = df_wpox[df_wpox_vax_mask]  # children with pox/number of vaxes
    
    df_nopox_mask = df_cpox['HAD_CPOX'] == 2  
    df_nopox = df_cpox[df_nopox_mask]  # children without pox
    df_nopox_vax_mask = df_nopox['P_NUMVRC'] >= 0
    df_nopox_vax = df_nopox[df_nopox_vax_mask]  # childred w/out pox/number of vaxes
      
    had_pox = list(df_wpox_vax['HAD_CPOX'])
    had_pox.extend(list(df_nopox_vax['HAD_CPOX']))
    
    num_vax = list(df_wpox_vax['P_NUMVRC'])
    num_vax.extend(list(df_nopox_vax['P_NUMVRC']))
    
    had_pox_s = pd.Series(had_pox)
    df_had_pox = pd.DataFrame(had_pox_s)
    df_had_pox = df_had_pox.rename(columns = {0: "had_chickenpox_column"})
                              
    num_vax_s = pd.Series(num_vax)
    df_num_vax = pd.DataFrame(num_vax_s)
    df_num_vax = df_num_vax.rename(columns = {0: "num_chickenpox_vaccine_column"})
    
    pox_vac_df = pd.concat([df_had_pox, df_num_vax], axis=1)
    pox_vac_df.sort_index(inplace=True)
    
    corr, pval=stats.pearsonr(pox_vac_df["had_chickenpox_column"],pox_vac_df["num_chickenpox_vaccine_column"])
    
    #print(pval)
    #print(corr)
    
    return corr
    raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

Output:
```
0.07044873460147982
```